rule_files:
  - /etc/prometheus/alerts.yml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'gps_api_latency_seconds_bucket{le="0.25",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+50x6'
      - series: 'gps_api_latency_seconds_bucket{le="0.5",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+80x6'
      - series: 'gps_api_latency_seconds_bucket{le="1.0",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+120x6'
      - series: 'gps_api_latency_seconds_bucket{le="2.5",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+140x6'
      - series: 'gps_api_latency_seconds_bucket{le="5.0",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+160x6'
      - series: 'gps_api_latency_seconds_bucket{le="10.0",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+180x6'
      - series: 'gps_api_latency_seconds_bucket{le="+Inf",path="/ingest/http",method="POST",status_code="200"}'
        values: '0+200x6'
    alert_rule_test:
      - eval_time: 6m
        alertname: ApiLatencyP95High
        exp_alerts:
          - exp_labels:
              alertname: ApiLatencyP95High
              path: /ingest/http
              severity: page
            exp_annotations:
              summary: "Latencia API p95 elevada"
  - interval: 1m
    input_series:
      - series: 'gps_ingestions_total{device_id="alpha"}'
        values: '0 0 0 0 0 0'
      - series: 'gps_ingestions_total{device_id="beta"}'
        values: '10 10 10 10 10 10'
    alert_rule_test:
      - eval_time: 6m
        alertname: IngestionTrafficDrop
        exp_alerts:
          - exp_labels:
              alertname: IngestionTrafficDrop
              severity: page
            exp_annotations:
              summary: "Ingestiones cayendo por debajo de 1/min"
  - interval: 1m
    input_series:
      - series: 'gps_device_last_seen_epoch{device_id="offline-1"}'
        values: '0 0 0 0 0 0'
    alert_rule_test:
      - eval_time: 6m
        alertname: DeviceOffline
        exp_alerts:
          - exp_labels:
              alertname: DeviceOffline
              device_id: offline-1
              severity: ticket
            exp_annotations:
              summary: "Dispositivo sin heartbeat >5m"
